<div class = "alert alert-warning" align="center"> 
    <img src="https://www.egd.mg/wp-content/uploads/2020/01/cropped-logo-home-110px.png" align="left" width="290"  margin="50px"> 
    <center><h1>Lycée français de Tananarive</h1> 
        <h2> Jupyter Hub,  Spécialité NSI </h2></center>
</div>

Algorithme de Recherche dichotomique
==
> Programmer une fonction `recherche_dico(L,n)` qui teste la présence d'un entier `n` dans une liste d'entiers ordonnés.

<!-- grade_id: cell-496c29b165e1e313, points: 0, grade: false, solution: true -->
```python
L = [i*2 for i in range(1000)]
def recherche_dico(T,n) :
    """ Teste la présence de n dant T
    Prec : T une liste d'entiers positifs ou nul triée par ordre croissant
           n un entier dont il faut tester la présence dans T
    Post : Un booléen, True si n est dans T, False sinon"""
    #YOUR CODE HERE
    raise NotImpl
```

<!-- grade_id: cell-d311ef0380de4f7f, points: 3, grade: true, solution: false -->
```python
L = [i*2 for i in range(1000)]
assert recherche_dico(L,900) == True
assert recherche_dico(L,57) == False
assert recherche_dico([1],2) == False
assert recherche_dico([1],1) == True
assert recherche_dico([1,500],50) == False
```

# Coût de la recherche dichotomique
**Dans quel cas la méthode de recherche dichotomique exécutera-t-elle le plus d'opérations élémentaires ?**

<!-- grade_id: cell-faeaaa10d7fadc2b, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

```python
"Dans le cas ou l'élément à rechercher se trouve dans la plus pire des positions!!!"
```

Ecrivons une fonction qui permette d'estimer le temps d'exécution de la fonction `recherche_dico(T,n)` en fonction du nombre de données du tableau :

```python
from time import process_time
def mesure_temps(n) :
    L = [i*2 for i in range(n)]
    m = (n-1)*2-1
    t1_start = process_time()  
    resultat = recherche_dico(L,m)
    t1_stop = process_time() 
    return (t1_stop-t1_start)  

```

Et voyons quelle est l'évolution du temps de calcul en fonction du nombre de données, et soyez un peu patien pour l'éxécution de de programme ! :

```python
import matplotlib.pyplot as plt

x = [i for i in range(0,1000000,100000)]
y=[]
for i in x :
    total = 0
    for j in range(20) :
        total = total + mesure_temps(i)
    y.append(total/20)

plt.plot(x,y)
```

**Justifiez le fait que le coût en temps ne soit pas linéaire. Ce coût vous parait-il plus faible ou plus grand qu'un coût en temps linéaire ?** 

<!-- grade_id: cell-341e1f571f68c84b, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Comparaison avec un recherche par parcours séquentiel de liste.
Les fonctions définies ci-dessous redéfinissent la fonction `appartient(T,n)` programmée lors d'un cours précédent et compare son cout d'éxécution à une recherche dichotomique

```python
def appartient(T,n) : 
    """Teste l'appartenance d'un élément à une liste
    Prec : Une liste T non vide, d'entiers positifs ou nuls
    Post : Un booléen True si n appartient à T, False si n n'appartient pas à T"""
    i = 0
    while i != len(T) and T[i] != n :
        i += 1
    if i == len(T) :
        return False
    else :
        return True
```

```python
def mesure_temps2(n) :
    L = [i*2 for i in range(n)]
    if (n//2)%2 == 0 :
        m= n/2+1
    else :
        m = n//2
    t1_start = process_time()  
    resultat = appartient(L,m)
    t1_stop = process_time() 
    return (t1_stop-t1_start) 
```

```python
x = [i for i in range(0,300,2)]
y1,y2=[],[]
for i in x :
    total1,total2 = 0,0
    for j in range(100) :
        total1 = total1 + mesure_temps(i)
    for j in range(100) :
        total2 = total2 + mesure_temps2(i)
    y1.append(total1/100)
    y2.append(total2/100)

plt.plot(x,y1,label = 'recherche dichotomique')
plt.plot(x,y2,label = 'recherche séquentielle')
plt.legend()
```

**Que montre ce graphe ?**

<!-- grade_id: cell-174138b5e1d67796, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Cependant dans cette comparaison, quelle particularité de la recherche dichotomique n'a pas été prise en compte dans sa comparaison avec le tri séquentiel ?**

<!-- grade_id: cell-b106ea2a9ae7a7aa, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE