<div class = "alert alert-warning" align="center"> 
    <img src="https://www.egd.mg/wp-content/uploads/2020/01/cropped-logo-home-110px.png" align="left" width="290"  margin="50px"> 
    <center><h1>Lycée français de Tananarive</h1> 
        <h2> Jupyter Hub,  Spécialité NSI </h2></center>
</div>

# Objectif de l'activité.
**L'objectif de cette activité est de déterminer si un codon d'un brin ARN appartient à une séquence d'ARN, et finalement de compter le nombre de codons identiques contenus dans le brin ARN**

# Créer un bout d'ARN !
## Le codon d'ARN.

<img src="RNA-codon.png" align="right"> 
Un codon est une séquence de **trois nucléotides** sur un acide ribonucléique messager (ARNm) spécifiant l'un des 22 acides aminés protéinogènes dont la succession sur l'ARN messager détermine la structure primaire de la protéine à synthétiser. Les ARN messagers étant constitués d'une succession de plusieurs dizaines à plusieurs centaines de nucléotides, ce sont également plusieurs dizaines à plusieurs centaines d'acides aminés qui peuvent être ainsi assemblés linéairement pour former des chaînes protéiques. **Quatre bases nucléiques** déterminent la séquence d'un ARN messager — adénine(A), uracile(U), guanine(G), et cytosine(C) — de sorte qu'il existe 43 = 64 codons différents

> Ecrire une fonction `codon()` qui renvoie un codon aléatoire de trois bases nucléiques codées sous la forme d'une chaine de caractéres, 'UUG' par exemple

On rappelle que python utilise la fonction `sorted(L)` pour trier une liste, et que dans le cas d'une chaine de caractéres ceux-ci sont trier par leur code ASCII code. 

`ACA` sera donc placé avant `ACC` car :

- 'A' : 65 
- 'C' : 67
- 'G' : 71
- 'U' : 125


<!-- grade_id: cell-e16a10b8cf65c66d, points: 0, grade: false, solution: true -->
```python
from random import randint

def codon() :
    """ Renvoie une chaine aléatoire de 3 carctéres qui représente un codon de l'ARN
    Prec : aucune, ne prend pas de paremètres.
    Post : renvoie une chaine aléatoire de 3 caractéres parmi A,U,G,C"""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-84712ba5f9ae87e8, points: 1, grade: true, solution: false -->
```python
assert len(codon()) == 3
C = codon()
assert 'A' in  C or 'U' in C or 'G' in C or 'C' in C == True
assert not 'B' in C == True
```

> Ecrire une fonction `arn(n)` qui renvoie sous forme de liste une branche d'ARN codée sous la forme d'une liste de codon 

Exemple de sortie pour `n=3` : ['UGG', 'ACA', 'UUA'] 

<!-- grade_id: cell-8cedbccbe3e63cf8, points: 0, grade: false, solution: true -->
```python
def arn(n) :
    """Renvoie une liste triée de codons sous forme d'une liste
    Prec : un entier qui représente le nombre de codon que contiendra la liste 
    Post : Une liste aléatoire de n chaines de trois caractéres crées par codon()"""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-6d10040f5de3f50c, points: 1, grade: true, solution: false -->
```python
assert len(arn(10)) == 10
A = arn(5)
assert ('A' in  A[0] or' U' in A[0] or 'G' in A[0] or 'C' in A[0]) == True
assert len(A[1]) == 3
assert (not 'H' in A[2]) == True
```

# Rechercher un codon dans un brin ARN
## Algorithme de recherche dichotomique.
> Ecrire une fonction `recherche(T,n)` en utilisant une méthode de recherche dichotomique d'un codon dans un brin ARN.

<!-- grade_id: cell-cd1c23ca495e5852, points: 0, grade: false, solution: true -->
```python
def recherche_dico(T,n) :
    """ Teste la présence de n dans T
    Prec : T une liste de codon d'au moins 2 éléments triée par ordre croissant
           n un codon sous la forme d'une chaine de 3 caractères.
    Post : Un booléen, True si n est dans T, False sinon"""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-483733432bf5294e, points: 3, grade: true, solution: false -->
```python
assert recherche_dico(arn(10),'HHH') == False
assert recherche_dico(['AAA','AGA'], 'AAA') == True
assert recherche_dico(['AAA','AGA'], 'AUU') == False
assert recherche_dico(['AAA','AGA','UUU'], 'UUU') == True
```

## Algorithme de recherche séquentielle
On écrit cette fois-ci une fonction beaucoup plus simple qui réalise elle une recherche séquentielle d'un codon dans la liste représentative du brin ARN

<!-- grade_id: cell-8d4cd52eea995289, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Verifions tout d'abord que votre fonction de rcherche dichotomique est plus rapide que cette fonction de recherche séquentielle. Pour cela créons une fonction de mesure du temps d'exécution du temps des deux fonctions. La mesure sera effectuée sur une liste de longeur `n`. 

La fonction renvoie un tuple permettant de comparer les temps d'éxécution. Pour comparer les deux fonctions, nous nous appuierons sur un temps d'exécution maximum, avec une recherche qui dans les deux cas renvoie `False`.

```python
from time import process_time

def mesure_temps(n) :
    """ Mesure le temps d'exécution de recherche dichotomique et de recherche séquentielle
    Prec : un entier n représentant la longueur de la liste de recherche
    Post : Un tuple :
            le temps de recherche par une méthode dichotomique
            le temps de recherche par une méthode séquentielle"""
    L = arn(n)
    m = "HHH"
    t_dicho = 0 
    for i in range(10) :
        t1_start = process_time()  
        resultat = recherche_dico(L,m)
        t1_stop = process_time() 
        t_dicho += t1_stop-t1_start
    t_dicho = t_dicho/10
    t_seq = 0
    for i in range(10) :
        t2_start = process_time()  
        resultat = recherche_seq(L,m)
        t2_stop = process_time() 
        t_seq += t2_stop-t2_start
    t_seq = t_seq/10
    return t_dicho,t_seq
```

<!-- grade_id: cell-a2a457b256afb5f8, points: 2, grade: true, solution: false -->
```python
assert (mesure_temps(100000)[0]<mesure_temps(100000)[1]) == True
assert (mesure_temps(10000)[0]<mesure_temps(10000)[1]) == True
assert (mesure_temps(1000)[0]<mesure_temps(1000)[1]) == True
assert (mesure_temps(100)[0]<mesure_temps(100)[1]) == True
```

```python
import matplotlib.pyplot as plt

x = [i for i in range(2,10000,100)]
y1,y2=[],[]
for i in x :
    y1.append(mesure_temps(i)[0])
    y2.append(mesure_temps(i)[1])

plt.plot(x,y1,label="dicho")
plt.plot(x,y2,label='seq')
plt.legend()
```

# Comptage du nombre de codons identiques.
## Recherche dichotomique de l'indice d'un codon
> En vous appuyant sur votre fonction de recherche dichotomique, précédemment écrite, **écrire une fonction `recherche_indice(T,codon)` qui renvoie l'indice du codon passé en paramètre** 

<!-- grade_id: cell-7c02297c8a9bc396, points: 0, grade: false, solution: true -->
```python
def recherche_indice(T,n) :
    """ compte le nombre de codon n dans T
    Prec : T une liste de codons d'au moins 2 éléments triée par ordre croissant
           n un codon sous la forme d'une chaine de 3 caractères.
    Post : un entier représentant l'indice du codon passé en paramètre
           -1 si le codon recherché n'est pas dans T"""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-7978f66c6ef13c31, points: 2, grade: true, solution: false -->
```python
assert recherche_indice(['AAA','AAA'],'AAA') == 0
assert recherche_indice(['AAA','AAA'],'ACC') == -1
assert recherche_indice(['AAA','UUU'],'UUU') == 1
assert recherche_indice(['AAA','ACA','UUU'],'ACA') == 1
assert recherche_indice(['AAA','ACA','ACU','UUU'],'ACG') == -1
```

## Comptage du nombre de condons.
> Ecrire une fonction `nombre_dicho(L,n)` qui compte le nombre de codons `n` contenus dans la liste `L`

On pourra rechercher l'indice du codon passé en paramètre avec la fonction `recherche_indice(T,n)`, pour le supprimer de la liste en incrémantant un compteur jusqu'à ce que ce codon ne se trouve plus dans la liste.

<!-- grade_id: cell-840d4083034cfe49, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-95a0542a2a81df94, points: 2, grade: true, solution: false -->
```python
assert nombre_dico(['AAA','AAA'],'AAA') == 2
assert nombre_dico(['AAA','AAA'],'CAC') == 0
assert nombre_dico(['AAA','AAC','AAG'],'AAA') == 1
assert nombre_dico(['AAA','AAA','AAG'],'AAA') == 2
assert nombre_dico(['AAA','AAG','AAG'],'AAG') == 2
```

On défini maintenant une fonction `nombre_seq(L,n)` qui compte de façon séquentielle ne nombre d'occurences du codon `n`

```python
def nombre_seq(L,n) :
    nb = 0 
    for i in L :
        if i == n :
            nb = nb +1
    return nb
```

Vérifions tout d'abord que les deux fonctions renvoient bien le même nombre :

```python
L_test = arn(1000)
nb1 = nombre_seq (L_test,'AAA')
nb2 = nombre_dico (L_test,'AAA')
assert nb1 == nb2
nb1 = nombre_seq (L_test,'CAG') 
nb2 = nombre_dico (L_test,'CAG')
assert nb1 == nb2
```

## Temps d'exécution.
Ecrivons maintenant une fonction qui compare les temps d'exécution des type de comptage, en utilisant une méthode dichotomique et une méthode séquentielle.

```python
def mesure_temps2(n) :
    """ Mesure le temps d'exécution de comptage dichotomique et de comptage séquentiel
    Prec : un entier n représentant la longueur de la liste de recherche
    Post : Un tuple :
            le temps de comptage par une méthode dichotomique
            le temps de comptage par une méthode séquentielle"""
    L = arn(n)
    m = "UUU"
    t_seq = 0
    for i in range(10) :
        t2_start = process_time()  
        resultat = nombre_seq(L,m)
        t2_stop = process_time() 
        t_seq += t2_stop-t2_start
    t_seq = t_seq/10
    t_dicho = 0 
    for i in range(10) :
        t1_start = process_time()  
        resultat = nombre_dico(L,m)
        t1_stop = process_time() 
        t_dicho += t1_stop-t1_start
    t_dicho = t_dicho/10
    return t_dicho,t_seq
```

```python
x = [i for i in range(2,100000,5000)]
y1,y2=[],[]
for i in x :
    L = mesure_temps2(i)
    y1.append(L[0])
    y2.append(L[1])

plt.plot(x,y1,label="dicho")
plt.plot(x,y2,label='seq')
plt.legend()
```

## Mieux que Python !
Python dispose d'une méthode que l'on peut appliquer sur les listes `L.count(n)` qui compte le nombre d'occurences de `n` dans `L`.

```python
def nombre_python(L,n) :
    return L.count(n)
```

Vérifions que nos 3 fonctions de comptage sont équivalentes :

```python
L_test = arn(1000)
assert nombre_seq(L,'AAA') == nombre_python(L,'AAA')
assert nombre_seq(L,'AHA') == nombre_python(L,'AHA')
assert nombre_seq(L,'CAG') == nombre_python(L,'CAG')
```

Et écrivons une fonction `mesure_temps3(n)` qui renvoie le temps d'exécution pour les trois fonctions de comptage, pour un nombre `n` passé en paramètre.

```python
def mesure_temps3(n) :
    L = arn(n)
    m = "UUU"
    t_seq = 0
    for i in range(10) :
        t2_start = process_time()  
        resultat = nombre_seq(L,m)
        t2_stop = process_time() 
        t_seq += t2_stop-t2_start
    t_seq = t_seq/10
    t_dicho = 0 
    for i in range(10) :
        t1_start = process_time()  
        resultat = nombre_dico(L,m)
        t1_stop = process_time() 
        t_dicho += t1_stop-t1_start
    t_dicho = t_dicho/10
    L = arn(n)
    t_python = 0
    for i in range(10) :
        t3_start = process_time()  
        resultat = nombre_python(L,m)
        t3_stop = process_time() 
        t_python += t3_stop-t3_start
    t_python = t_python/10
    return t_python,t_dicho,t_seq
```

```python
x = [i for i in range(2,1000000,500000)]
y1,y2,y3=[],[],[]
for i in x :
    L = mesure_temps3(i)
    y1.append(L[0])
    y2.append(L[1])
    y3.append(L[2])

plt.plot(x,y1,label="python")
plt.plot(x,y2,label='dicho')
plt.plot(x,y3,label='seq')
plt.legend()
```

**En conclusion : Parfois pour gagner du temps il faut accepter d'en perdre...**

**Visiblement les trois fonctions de recherches ont une complexité en $\mathcal{O}(n)$, essayez de justifier cette affirmation.**

<!-- grade_id: cell-49febc41051451f1, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```